# Python Coding Exercises
Complete each of the unfinished coding exercises below.  You can check all of your answers by clicking on the "Validate" button at the top of this window.  If you don't see a validate button you need to activate nbgrader.

## Type Hints
Note that the examples below make use of Python *type hints* to document expected argument and return types. For example, the first function declaration below could have been written as follows:

```python
def in_interval(x, a, b):
    """ Return True if the number x is in the interval [a, b)"""
```

Instead, we've written:

```python
def in_interval(x: float, a: float, b: float) -> bool:
    """ Return True if the number x is in the interval [a, b)"""
```

This is not a syntactic requirement in Python, and the Python interpreter will not enforce type hints by default. We probably won't use type hints consistently this semester, but they can be helpful when developing larger-scale Python programs.

Follow the link below for more information on type hints:

<https://docs.python.org/3.8/library/typing.html>

## Name & Honor Code

Enter your name in the cell below.  By doing so, you certify that this submission conforms to the JMU Honor code.

<!-- grade_id: Names, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

***
## QUESTION 1

<!-- grade_id: Q1, points: 0, grade: false, solution: true -->
```python
# QUESTION 1 - Simple Logic and Arithmetic

def in_interval(x: float, a: float, b: float) -> bool:
    """ Return True if the number x is in the interval [a, b)"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: Q1Tests, points: 1, grade: true, solution: false -->
```python
# TESTS FOR QUESTION 1

assert in_interval(0.9, 1.0, 2.0) is False
assert in_interval(1.0, 1.0, 2.0) is True
assert in_interval(1.5, 1.0, 2.0) is True
assert in_interval(2.0, 1.0, 2.0) is False
assert in_interval(2.1, 1.0, 2.0) is False
assert in_interval(0.5, 0.5, 0.5) is False
assert in_interval(0.5, 0.5, 0.6) is True
```

***
## QUESTION 2

<!-- grade_id: Q2, points: 0, grade: false, solution: true -->
```python
# Question 2 - Simple Lists and Strings

from typing import List

def string_lengths(strings: List[str]) -> List[int]:
    """ Return a list containing the lengths of the strings in the provided list.
        
        >>> string_lengths(["House", "Car", "Tree"])
        [5, 3, 4]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: Q2Tests, points: 1, grade: true, solution: false -->
```python
# TESTS FOR QUESTION 2

assert string_lengths(["House", "Car", "Tree"]) == [5, 3, 4]
assert string_lengths(["House", "Car", "Tree", "A"]) == [5, 3, 4, 1]
assert string_lengths(["Bar"]) == [3]
```

***
## QUESTION 3

<!-- grade_id: cell-2400bdc2a79de90f, points: 0, grade: false, solution: true -->
```python
# Question 3 - Dictionaries

from typing import Dict, List

def create_index(strings: List[str]) -> Dict[str, List[int]]:
    """ Return a dictionary mapping from each distinct string in the provided list to 
    an ordered list containing the positions where it appears.
    
    >>> create_index(["A", "B", "A", "C"])
    {"A": [0, 2], "B": [1], "C": [3]}
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()


```

<!-- grade_id: cell-30b1d4140c779238, points: 1, grade: true, solution: false -->
```python
# TESTS FOR QUESTION 3

assert create_index(["A"]) == {"A": [0]}
assert create_index(["A", "B", "A", "C"]) == {"A": [0, 2], "B": [1], "C": [3]}
assert create_index(["A", "A", "A", "A"]) == {"A": [0, 1, 2, 3]}
assert create_index(["House", "Car", "Cat", "Car"]) == {"House": [0], "Car": [1, 3], "Cat": [2]}
```

***
## QUESTION 4

There are many ways to solve this problem.  One approach is to store the current k-nearest
numbers as distance/number tuples and then use the Python list sort method to sort the tuples.  By default, Python will sort first by the first entry in the tuple, then by the second. For example:

```python
tuples = [(3, 1), (1, 7), (1, 2), (0, 12)]
tuples.sort()
print(tuples)
```

<!-- grade_id: cell-64fe2dd71c8ff278, points: 0, grade: false, solution: true -->
```python
from typing import List

def k_nearest(x: float, numbers: List[float], k: int=1) -> List[float]:
    """ Return a list containing the k nearest numbers to x.  The returned
    list must be ordered by proximity to x, with the closest number appearing
    first.  If the length of the list of numbers is less than or equal to k, 
    then all numbers will be returned.
    
    >>> k_nearest(3.0, [3.0, 3.7, 27.5, -22.7, 2.9], 3)
    [3.0, 2.9, 3.7]
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-3a12182dc6a41fe0, points: 1, grade: true, solution: false -->
```python
# TESTS FOR QUESTION 4

points = [3.0, 3.7, 27.5, -22.7, 2.9]

assert k_nearest(3.0, points, 3) == [3.0, 2.9, 3.7]
assert k_nearest(3.0, points) == [3.0]
assert k_nearest(3.0, points, 2) == [3.0, 2.9]
assert k_nearest(-30.0, points, 3) == [-22.7, 2.9, 3.0]
assert k_nearest(1000.0, points, 100) == [27.5, 3.7, 3.0, 2.9, -22.7]

```